# Assignment 1
For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web.



Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"
    
    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    #raise NotImplementedError()
```

## Part A

Find a list of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re
def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""
    
    #Method 2: All the names start with capital letters. Find all words in the string starting with capital letters
    #[A-Z] means look for words that start with capital letters
    #[\w] means that any letter can follow the capital letter
    # * tells us that any letters can follow the capital letter
    
    print(re.findall("[A-Z][\w]*", simple_string))
    print(len(re.findall("[A-Z][\w]*", simple_string)))
    
    return re.findall("[A-Z][\w]*", simple_string)
    
    #Method 1: look for Amy, Mary, Ruth, and Peter and concatenate them together (doesn't really use regex)
    #solution = re.findall("Amy", simple_string) + re.findall("Mary", simple_string) + re.findall("Ruth", simple_string) + re.findall("Peter", simple_string)
    #print(len(solution))
    #return solution
    #raise NotImplementedError()

```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"

```

Output:
```
['Amy', 'Mary', 'Ruth', 'Peter']
4
```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in 
a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re
def grades():
    with open ("assets/grades.txt", "r") as file:
        grades = file.read()
        #print(grades) looks like it's Firstname Lastname: Grade, grade can be 'A', 'B', or 'C'.
               
        # Method:
        #1) do a findall() to find a list of students with B grades after the colon.
        #just keep names in the group. colon space B doesn't need to be added to the list so no () for them
        
        print(re.findall("([\w]*[\s][\w]*)\:[\s][B]", grades))
        return re.findall("([\w]*[\s][\w]*)\:[\s][B]", grades)
        
       

    #raise NotImplementedError()
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(grades()) == 16

```

Output:
```
['Bell Kassulke', 'Simon Loidl', 'Elias Jovanovic', 'Hakim Botros', 'Emilie Lorentsen', 'Jake Wood', 'Fatemeh Akhtar', 'Kim Weston', 'Yasmin Dar', 'Viswamitra Upandhye', 'Killian Kaufman', 'Elwood Page', 'Elodie Booker', 'Adnan Chen', 'Hank Spinka', 'Hannah Bayer']
```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:
* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re
def logs():
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()
    
    #print(logdata)
    
    #Method:
    #1) figure out the regex to get what we need
    #2) going to need finditer() to return a dictionary. don't give a key to stuff we don't care about, like '-'.
        #2a) create a pattern to fit everything, and call the titles host, user_name, time, and request
        #2b) make the for loop to iterate through logdata.txt with finditer() to throw everything into the dictionary
    #3) return the dictionary

    #1) figure out the regex to get what we need
        #figure out regex for IP addresses
        #print(re.findall("([\w]*[.][\w]*[.][\w]*[.][\w]*)", logdata))
    
        #figure out regex for usernames
        #([\s]\-[\s]) is throwaway, while we want either a username of letters/digits or a '-' symbol: ([\w\-]*). throwaway space after that too
        #print(re.findall("([\s]\-[\s])([\w\-]*)([\s])", logdata))
        #put it together
        #print(re.findall("([\w]*[.][\w]*[.][\w]*[.][\w]*)([\s]\-[\s])([\w\-]*)([\s])", logdata))
    
        #figure our regex for times
        #[ w / w / w : w : w : w s - w ] s   based on the example, it looks like we don't care about sq brackets either
        #print(re.findall("(\[)([\w]*\/[\w]*\/[\w]*\:[\w]*\:[\w]*\:[\w]*[\s]\-[\w]*)(\])", logdata))
        #put it together
        #print(re.findall("([\w]*[.][\w]*[.][\w]*[.][\w]*)([\s]\-[\s])([\w\-]*)([\s])(\[)([\w]*\/[\w]*\/[\w]*\:[\w]*\:[\w]*\:[\w]*[\s]\-[\w]*)(\])", logdata))

        #figure out regex for requests
        # need to get rid of the double quotes. there's random stuff at the back too, to get rid of
        #" requestType everythingThatIsntADoubleQuote" " s w s w    " s w s w is random stuff we don't need
        #print(re.findall("([\s]\")([\w]*[^\"]*)(\"[\s][\w]*[\s][\w]*)", logdata))
        #put it together
        #print(re.findall("([\w]*[.][\w]*[.][\w]*[.][\w]*)([\s]\-[\s])([\w\-]*)([\s])(\[)([\w]*\/[\w]*\/[\w]*\:[\w]*\:[\w]*\:[\w]*[\s]\-[\w]*)(\])([\s]\")([\w]*[^\"]*)(\"[\s][\w]*[\s][\w]*)", logdata))
    
    #2) going to need finditer() to create a dictionary. don't give a key to stuff we don't care about, like '-'.
        #2a) create a pattern to fit everything, and call the titles host, user_name, time, and request
        pattern = """
        (?P<host>([\w]*[.][\w]*[.][\w]*[.][\w]*))     # IP address
        ([\s]\-[\s])       # don't care
        (?P<user_name>([\w\-]*))    # username
        ([\s])    #don't care
        (\[)    #don't care about square bracket open
        (?P<time>([\w]*\/[\w]*\/[\w]*\:[\w]*\:[\w]*\:[\w]*[\s]\-[\w]*))    #timestamp
        (\])    #don't care about square bracket close
        ([\s]\")    #don't care about space and double quote open
        (?P<request>([\w]*[^\"]*))    #request
        (\"[\s][\w]*[\s][\w]*)    #don't care about double quote close and numbers after that
        """      
        
        #2b) make the for loop to iterate through logdata.txt with finditer() to throw everything into a dictionary
        solution = []
        for item in re.finditer(pattern, logdata, re.VERBOSE):
            solution.append(item.groupdict())

    #3) return the dictionary
    #print(solution)
    return solution



```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

```

```python

```

```python

```